```python
import otter
grader = otter.Notebook("proj3.ipynb")
```

```python
import numpy as np
import pandas as pd
```

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
!psql postgresql://jovyan@127.0.0.1:5432/ucb_buildings -c 'SELECT pg_terminate_backend(pg_stat_activity.pid) FROM pg_stat_activity WHERE datname = current_database()  AND pid <> pg_backend_pid();'
!psql postgresql://jovyan@127.0.0.1:5432/postgres -c 'DROP DATABASE IF EXISTS ucb_buildings'
!psql postgresql://jovyan@127.0.0.1:5432/postgres -c 'CREATE DATABASE ucb_buildings'
!psql postgresql://jovyan@127.0.0.1:5432/ucb_buildings -f data/buildings.sql -q
```

Output:
```
pg_terminate_backend 
----------------------
 t
(1 row)

DROP DATABASE
CREATE DATABASE
 set_config 
------------
 
(1 row)
```

```python
from data101_utils import GradingUtil
grading_util = GradingUtil("proj3")
grading_util.prepare_autograder("ucb_buildings")
```

Output:
```
Successfully opened database connection
```

```python
grading_util.run_sql("SELECT * FROM data LIMIT 10;")
```

Output:
```
time                                    id     value
0 2018-06-12 23:00:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65182.74
1 2018-06-12 22:45:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65182.51
2 2018-06-12 22:30:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65182.27
3 2018-06-12 22:15:10+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65182.02
4 2018-06-12 22:00:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65181.77
5 2018-06-12 21:45:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65181.54
6 2018-06-12 21:30:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65181.30
7 2018-06-12 21:15:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65181.04
8 2018-06-12 21:00:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65180.80
9 2018-06-12 20:45:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06  65180.57
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


From looking at the schema itself  I see there's some tables that describe the metadata of something or the mapping of two objects; this optimizes both querying operations and storage, referecing the schema now: real_estate_metadata and metadata will cause for an easy lookup of attributes causing less repetition of the same values in multiple tables, therefore optimizing storage.  And mapping describe existing relationships like buildings_site_mapping and mapping; which it's easy for the user to see how the data interacts with each other which will result in the optimization of lookups. 

```python
--DROP VIEW IF EXISTS view_one;
WITH view_one AS ( 
SELECT DISTINCT buildings_site_mapping.building
FROM buildings_site_mapping
    )

SELECT view_one.building, JSON_AGG(real_estate_metadata)
FROM real_estate_metadata
JOIN view_one ON view_one.building = real_estate_metadata.building_name
GROUP BY view_one.building
HAVING COUNT(real_estate_metadata.building_name) > 1
ORDER BY view_one.building ASC
```

```python
--result_1b
SELECT *
FROM real_estate_metadata
WHERE building_name = 'ALUMNI HOUSE';

```

Output:
```
location building building_name          address      city_name  \
0       BERKELEY     1215  ALUMNI HOUSE  CORE C22A2M0PUS       BERKELEY   
1         IRVINE     9207  ALUMNI HOUSE      CORE CAMPUS         IRVINE   
2  FRANCISC SOAN     2032  ALUMNI HOUSE   PARNASS US AVE  FRANCISC OSAN   

          county        category   osfg  bldg_asf  levels  year owner mp_code  \
0        ALAMEDA         GENERAL  15590      8719       2  1954    UC       P   
1         ORANGE         GENERAL   4027      2549       1  1984    UC       P   
2  FRANCISC OSAN  HEALTH SCIENCE   7217      5079       3  1915    UP       P   

   book_value  
0    106819.0  
1     41981.0  
2    135923.0
```

```python
result_1b = grading_util.run_file("1b")
result_1b
```

Output:
```
building  \
0  ALUMNI HOUSE   
1      CAMPBELL   
2         HERTZ   
3         SOUTH   
4        SPROUL   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                    json_agg
```

```python

WITH view_one AS  (
SELECT DISTINCT real_estate_metadata.building_name
FROM real_estate_metadata
)


SELECT view_one.building_name, JSON_AGG(buildings_site_mapping)
FROM buildings_site_mapping
JOIN view_one ON view_one.building_name = buildings_site_mapping.building
GROUP BY view_one.building_name
HAVING COUNT(buildings_site_mapping.building) > 1
ORDER BY view_one.building_name ASC
```

```python
result_1c = grading_util.run_file("1c")
result_1c
```

Output:
```
building_name  \
0   EDWARDS FLD   
1      FAC CLUB   
2  HAAS STU BLD   
3         SIMON   

                                                                                                                                          json_agg  
0                       [{'site': 'Edwards Stadium East', 'building': 'EDWARDS FLD'}, {'site': 'Edwards Stadium West', 'building': 'EDWARDS FLD'}]  
1                                [{'site': 'Men's Faculty Club', 'building': 'FAC CLUB'}, {'site': 'Womens Faculty Club', 'building': 'FAC CLUB'}]  
2  [{'site': 'Haas School Student Services (Water)', 'building': 'HAAS STU BLD'}, {'site': 'Haas School of Business', 'building': 'HAAS STU BLD'}]  
3                                         [{'site': 'Simpson Center', 'building': 'SIMON'}, {'site': 'Boalt and Simon Hall', 'building': 'SIMON'}]
```

Things to note: 
1. From the insights in 1b we can't determine a unique building in the sensor data. Sensor data's attributes are value, id and time, while real_estate_metadata tells information about the building itself and none of those attributes have any of the sensor data.
2. It's important to note that  buildings_site_mapping.building couldn' be a foreign key poiting to real_estate_metadata.building_name. As there's multiple repetitions real_estate_metadata.building_name that correspond to a buildings_site_mapping.building entry therefore breaking referential integrity is it has a one-to-many relationship.


### Outlier Detection


```python
--result_2a

WITH median_values AS (
SELECT id, 
PERCENTILE_DISC(0.5) WITHIN GROUP (ORDER BY value) AS median
FROM data
GROUP BY id
), 

two AS ( 
SELECT one.time, one.id, one.value, abs (one.value - view.median) AS premad
FROM data one
JOIN median_values view ON view.id = one.id
),

mad_table AS (
SELECT id,
PERCENTILE_DISC(0.5) WITHIN GROUP (ORDER BY premad) AS mad
FROM two
GROUP BY id
)

 
SELECT 
two.time, 
two.id, 
two.value, 
median_values.median, 
mad_table.mad, 
CASE 
    WHEN ABS(two.value - median_values.median) >  3 * 1.4826 * mad_table.mad
    THEN TRUE 
    WHEN mad_table.mad = 0 THEN FALSE 
    ELSE FALSE 
    END AS is_outlier 

FROM two
JOIN mad_table ON mad_table.id = two.id
JOIN median_values ON median_values.id = two.id
;

```

Output:
```
time                                    id  \
0     2018-06-12 23:00:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06   
1     2018-06-12 22:45:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06   
2     2018-06-12 22:30:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06   
3     2018-06-12 22:15:10+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06   
4     2018-06-12 22:00:09+00:00  a3d3326f-20ab-5f1d-97c7-f3084df43f06   
...                         ...                                   ...   
99995 2018-06-11 02:15:00+00:00  a8dc3e90-6bb7-56e4-bf10-bf07694da31b   
99996 2018-06-11 02:00:00+00:00  a8dc3e90-6bb7-56e4-bf10-bf07694da31b   
99997 2018-06-11 01:45:00+00:00  a8dc3e90-6bb7-56e4-bf10-bf07694da31b   
99998 2018-06-11 01:30:00+00:00  a8dc3e90-6bb7-56e4-bf10-bf07694da31b   
99999 2018-06-11 01:15:00+00:00  a8dc3e90-6bb7-56e4-bf10-bf07694da31b   

           value     median     mad  is_outlier  
0      65182.740  65130.270  22.550       False  
1      65182.510  65130.270  22.550       False  

```

```python
result_2a = grading_util.run_file("2a")[1]
result_2a
```

Output:
```
time                                    id  value  \
0  2018-06-07 00:00:00+00:00  a481d4a8-44f6-5f63-b469-85571a018c6f  812.6   
1  2018-06-07 00:00:00+00:00  a51d3bed-a8b6-5321-96a1-c5542caca70c   75.0   
2  2018-06-07 00:00:00+00:00  a6336afc-cdb7-5f87-a4bf-67411ea1b5cc   54.0   
3  2018-06-07 00:00:00+00:00  a652779f-8433-5ad3-9310-f4f1062e4d04   14.0   
4  2018-06-07 00:00:00+00:00  a760a23a-8de4-5ce3-939a-826bf31f01df  109.0   
..                       ...                                   ...    ...   
95 2018-06-07 00:00:00+00:00  a7145dcc-90e5-5bfb-9b1c-277615779c16    4.0   
96 2018-06-07 00:00:00+00:00  a71d821d-d6e6-5ee9-86fa-2b075f13c2cc   99.0   
97 2018-06-07 00:00:00+00:00  a71fbdfd-26ad-562c-b193-9dc414b851f0  122.9   
98 2018-06-07 00:00:00+00:00  a744b292-f4cf-5f85-9e4a-5e626e19d5af  107.0   
99 2018-06-07 00:00:00+00:00  a749b51c-0be8-5b73-bf67-de452e0270ab   64.9   

     median     mad  is_outlier  
0   727.500  16.200        True  
1     6.186   4.957        True
```

<br><br>

---

### Outlier Handling (Winsorization)


```python
--result_2b

CREATE OR REPLACE VIEW cleaned_data AS

-- BEGIN SOLUTION
WITH median_values AS (
SELECT id, 
PERCENTILE_DISC(0.5) WITHIN GROUP (ORDER BY value) AS median
FROM data
GROUP BY id
), 

two AS ( 
SELECT one.time, one.id, one.value, abs (one.value - view.median) AS premad
FROM data one
JOIN median_values view ON view.id = one.id
),

mad_table AS (
SELECT id,
PERCENTILE_DISC(0.5) WITHIN GROUP (ORDER BY premad) AS mad
FROM two
GROUP BY id
)

 
SELECT 
two.time, 
two.id, 
two.value, 
median_values.median, 
mad_table.mad, 
CASE 
    WHEN mad_table.mad = 0 THEN FALSE
    WHEN ABS(two.value - median_values.median) >  3 * 1.4826 * mad_table.mad
    THEN TRUE 
    ELSE FALSE 
    END AS is_outlier,
    
    CASE
    WHEN mad_table.mad = 0 THEN two.value
    
    WHEN ABS(two.value - median_values.median) <  3 * 1.4826 * mad_table.mad THEN two.value

    WHEN two.value > median_values.median + 3 * 1.4826 * mad_table.mad
    THEN  median_values.median + 3 * 1.4826 * mad_table.mad

   WHEN two.value < median_values.median - 3 * 1.4826 * mad_table.mad
    THEN  median_values.median -  3 * 1.4826 * mad_table.mad
    ELSE two.value 
    END AS clean_value
    
    

FROM two
JOIN mad_table ON mad_table.id = two.id
JOIN median_values ON median_values.id = two.id
;
-- END SOLUTION

(SELECT * FROM cleaned_data WHERE is_outlier ORDER BY time, id LIMIT 50)
UNION ALL
(SELECT * FROM cleaned_data WHERE NOT is_outlier ORDER BY time, id LIMIT 50);

```

#### Entity Resolution (units + locations)

```python
result_2b = grading_util.run_file("2b")[1]
result_2b
```

Output:
```
time                                    id  value  \
0  2018-06-07 00:00:00+00:00  a481d4a8-44f6-5f63-b469-85571a018c6f  812.6   
1  2018-06-07 00:00:00+00:00  a51d3bed-a8b6-5321-96a1-c5542caca70c   75.0   
2  2018-06-07 00:00:00+00:00  a6336afc-cdb7-5f87-a4bf-67411ea1b5cc   54.0   
3  2018-06-07 00:00:00+00:00  a652779f-8433-5ad3-9310-f4f1062e4d04   14.0   
4  2018-06-07 00:00:00+00:00  a760a23a-8de4-5ce3-939a-826bf31f01df  109.0   
..                       ...                                   ...    ...   
95 2018-06-07 00:00:00+00:00  a7145dcc-90e5-5bfb-9b1c-277615779c16    4.0   
96 2018-06-07 00:00:00+00:00  a71d821d-d6e6-5ee9-86fa-2b075f13c2cc   99.0   
97 2018-06-07 00:00:00+00:00  a71fbdfd-26ad-562c-b193-9dc414b851f0  122.9   
98 2018-06-07 00:00:00+00:00  a744b292-f4cf-5f85-9e4a-5e626e19d5af  107.0   
99 2018-06-07 00:00:00+00:00  a749b51c-0be8-5b73-bf67-de452e0270ab   64.9   

     median     mad  is_outlier  clean_value  
0   727.500  16.200        True   799.554360  
1    
```

From inspection it seems there are inconsistensies, duplicates and ambiguity; for example in the last rows there's gallons and gpm (gallons per minute) which would create opportunities for cleaning the data as they might refer to the same thing, or mark a clear distinction to in case they are different in order to keep the integrity of the data.

```python
SELECT units FROM metadata;

```

Output:
```
units
0         kWh
1          kW
2          kW
3          kW
4          kW
...       ...
9504  Gallons
9505      Gpm
9506      Gpm
9507      Gpm
9508      Lbs

[9509 rows x 1 columns]
```

```python
--result_3b
SELECT COUNT(DISTINCT units) AS num_unique_units, 
COUNT(DISTINCT LOWER(units)) AS num_unique_units_ignore_case

FROM 
metadata

```

```python
result_3b = grading_util.run_file("3b")
result_3b
```

Output:
```
num_unique_units  num_unique_units_ignore_case
0                34                            29
```

```python
grading_util.run_sql("""
SELECT * FROM metadata;
""")
```

Output:
```
id  \
0     fb5267b9-8cf4-5710-979e-5ab7d617e6e5   
1     4e7e5872-a3f7-5300-8433-b4dfd0254cb1   
2     a2fbb537-3e9c-5c47-94c9-2b5cc3735f48   
3     4eb9d406-9aeb-56ab-aa9d-aaa368679e1f   
4     3c44bdad-d1a9-5155-b786-fb2ce6e1e76e   
...                                    ...   
9504  13e85c96-6104-53d0-b9d7-935e8f378e5c   
9505  924e8a05-8534-5b5a-a9a7-1ea4b1c7e14f   
9506  a6dc58d3-a002-5329-b528-f865228656d0   
9507  20174662-7781-55c4-bd98-4bf1c08c7fac   
9508  161e7db4-3c68-5c65-bd67-67f98f8087da   

                                       class                  site    units  
0                  Main Electric Meter (kWh)  2000 Carleton Street      kWh  
1            Main Electric Meter Demand (kW)  2000 Carleton Street       kW  
2               Main Electric Meter Min (kW)  2000 Carleton Street       kW  
3     Main Electric Meter Instantaneous (kW)  2000 Carleton Street       kW  
4               Main Electric Meter Max (kW)  2000 Carleton Street       kW  
...                
```



### Checking consistency of units



```python
--result_3c
WITH view_one AS (
    SELECT class, LOWER(units) as lower_caps
    FROM metadata
    GROUP BY class, lower_caps
)

SELECT 

CASE WHEN COUNT(DISTINCT lower_caps) = 1 THEN TRUE
ELSE FALSE END AS 
are_units_consistent
FROM view_one

```

```python
result_3c = grading_util.run_file("3c")
result_3c
```

Output:
```
are_units_consistent
0                 False
```

Some locations have ambiguous location names that are not specific enough. To provide an example, row 5273 has under location the string 'AG FIELD STAT' and in the address column it actually provide this example. Upon further investigation online, AG FIELD STAT doesn't throw meaning results as it references a government agency.

```python
SELECT * FROM uc_locations;
```

Output:
```
loc_id       loc_name
0        1  AG FIELD STAT
1        2       BERKELEY
2        3          DAVIS
3        4         IRVINE
4        5    LOS ANGELES
5        6         MERCED
6        7      RIVERSIDE
7        8      SAN DIEGO
8        9  SAN FRANCISCO
9       10  SANTA BARBARA
10      11     SANTA CRUZ
11      12     SYSTEMWIDE
```


---
### Used uc_locations to build a cleaned mapping:



```python
--result_3e
WITH ranked AS (
SELECT 

real_estate_metadata.building_name AS building_name,

real_estate_metadata.address AS address,

real_estate_metadata.location AS location,

uc_locations.loc_name AS clean_location,

--word_similarity(real_estate_metadata.location, uc_locations.loc_name) AS similarity_score,

RANK() OVER (PARTITION BY real_estate_metadata.location ORDER BY word_similarity(real_estate_metadata.location, uc_locations.loc_name) DESC) AS rank

FROM real_estate_metadata, uc_locations
),

ranked_two AS (
SELECT building_name, address, location, clean_location 
FROM ranked 
WHERE rank = 1

)
SELECT * 
FROM ranked_two;

```

```python
result_3e = grading_util.run_file("3e")
result_3e
```

Output:
```
building_name                address       location clean_location
0        SC ACAD STR        IRVINE B76L0V1D  AG FIELD STAT  AG FIELD STAT
1       WS 304 HOUSE     OAKLAND 1A7V3E53 W  AG FIELD STAT  AG FIELD STAT
2     H ADMIN HBL DG     UNIVERSIT Y40 R7D0  AG FIELD STAT  AG FIELD STAT
3      D POTTING SDH        HOLTON R10D04 E  AG FIELD STAT  AG FIELD STAT
4      H NUTRI BHARN     UNIVERSIT 4Y0 R7D0  AG FIELD STAT  AG FIELD STAT
...              ...                    ...            ...            ...
5271            UCDC     RHODE 1IS4L3A5N0 D    SYSTEMWI DE     SYSTEMWIDE
5272            UCPC       MERIDIAN PARKWAY    SYSTEMWI DE     SYSTEMWIDE
5273       RINCON 70       -BLAKE 1H1O1U1SE    SYSTEMWI DE     SYSTEMWIDE
5274    BERKWAY 2120  BERKELEY W70A RYINCON    SYSTEMWI DE     SYSTEMWIDE
5275        FRANKLIN        FRANKLIN STREET    SYSTEMWI DE     SYSTEMWIDE

[5276 rows x 4 columns]
```



## Part 4 – Interpolating Missing Data (gap-filling + linear interpolation)

Goal: Detect missing sensor readings in time, fabricate the missing timestamps, and interpolate reasonable values.


```python
--result_4a

CREATE OR REPLACE VIEW gaps AS

-- BEGIN SOLUTION
WITH diffs AS (
SELECT 
id, 
DATE_TRUNC('minute', time) AS time,
value, 
LAG(DATE_TRUNC('minute', time), 1) OVER (PARTITION BY id ORDER BY time) AS lag_time,
    
LAG(value, 1) OVER (PARTITION BY id ORDER BY time) AS lag_value, 
    
DATE_TRUNC('minute', time) - LAG(DATE_TRUNC('minute', time), 1) OVER (PARTITION BY id ORDER BY time) AS time_diff
FROM data

)

SELECT * 
FROM diffs
WHERE time_diff >= '30 minutes';
-- END SOLUTION

SELECT * FROM gaps ORDER BY id, time;
```

```python
result_4a = grading_util.run_file("4a")[1]
result_4a
```

Output:
```
id                      time    value  \
0    a3d47b1a-985e-5395-a6ee-719dad9b580f 2018-06-13 09:15:00+00:00    0.056   
1    a3e8e405-0eed-59b6-8747-fe892a6f93de 2018-06-09 04:15:00+00:00  280.800   
2    a46fb790-028c-5d17-a3e6-7d08daec0c03 2018-06-13 09:15:00+00:00  281.460   
3    a470ac6d-d448-522b-9b74-8ed27f24de25 2018-06-09 04:15:00+00:00    0.004   
4    a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 00:45:00+00:00   27.273   
..                                    ...                       ...      ...   
766  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 08:15:00+00:00    5.400   
767  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 13:00:00+00:00    5.400   
768  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 14:30:00+00:00    2.700   
769  a8d640b1-c17f-51cf-b06f-de6e704e87c1 2018-06-13 09:15:00+00:00  106.000   
770  aa6d2954-2a1f-5a86-9bb2-69ad7326789e 2018-06-13 09:15:00+00:00  286.700   

                     lag_time  lag_value       time_diff  
0   2018-06-13 08:4
```

### Filling in timestamps with GENERATE_SERIES

Used GENERATE_SERIES over each gap to manufacture new timestamps every 15 minutes between start and end.




```python
--result_4b 

CREATE OR REPLACE VIEW complete AS

-- BEGIN SOLUTION

WITH gaps AS (
SELECT 
id, 
DATE_TRUNC('minute', time) AS time,
value, 
LAG(DATE_TRUNC('minute', time), 1) OVER (PARTITION BY id ORDER BY time) AS lag_time,
    
LAG(value, 1) OVER (PARTITION BY id ORDER BY time) AS lag_value, 
    
DATE_TRUNC('minute', time) - LAG(DATE_TRUNC('minute', time), 1) OVER (PARTITION BY id ORDER BY time) AS time_diff
FROM data
),

gaps_two AS (
SELECT * 
FROM gaps
WHERE time_diff >= '30 minutes'
) 

SELECT 
id, time, value

FROM data 

UNION ALL

SELECT 

id, 

generated_time AS time,

NULL AS value 

FROM gaps_two

CROSS JOIN LATERAL GENERATE_SERIES(lag_time + '15 minutes'::INTERVAL, 
time - '15 minutes'::INTERVAL, '15 minutes'::INTERVAL) AS generated_time;
;
-- END SOLUTION

SELECT * FROM complete ORDER BY id, time LIMIT 100;
```

```python
grading_util.run_sql("""
SELECT *
FROM GENERATE_SERIES(1, 10, 3);
""")
```

Output:
```
generate_series
0                1
1                4
2                7
3               10
```

```python
result_4b = grading_util.run_file("4b")[1]
result_4b
```

Output:
```
id                      time     value
0   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:00:09+00:00  65085.99
1   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:15:09+00:00  65086.16
2   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:30:09+00:00  65086.35
3   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:45:09+00:00  65086.52
4   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 01:00:09+00:00  65086.71
..                                   ...                       ...       ...
95  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 23:45:09+00:00  65100.44
96  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:00:09+00:00  65100.63
97  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:15:09+00:00  65100.80
98  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:30:09+00:00  65100.97
99  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:45:09+00:00  65101.16

[100 rows x 3 columns]
```

```python
grading_util.run_sql("""
SELECT *
FROM complete;
""")
```

Output:
```
id                      time  \
0       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 23:00:09+00:00   
1       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:45:09+00:00   
2       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:30:09+00:00   
3       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:15:10+00:00   
4       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:00:09+00:00   
...                                      ...                       ...   
102381  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 13:45:00+00:00   
102382  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 14:00:00+00:00   
102383  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 14:15:00+00:00   
102384  a8d640b1-c17f-51cf-b06f-de6e704e87c1 2018-06-13 09:00:00+00:00   
102385  aa6d2954-2a1f-5a86-9bb2-69ad7326789e 2018-06-13 09:00:00+00:00   

           value  
0       65182.74  
1       65182.51  
2       65182.27  
3       65182.02  
4       65181.77  
...          ...  
102381       NaN  
1
```

### Question 4c:**Linear interpolation (forward / backward passes):**

```python
--result_4c
CREATE TABLE temp (
    id INTEGER,
    num INTEGER
);

INSERT INTO temp VALUES
    (1, 1),
    (2, 2),
    (3, NULL),
    (4, NULL),
    (5, 3),
    (6, NULL),
    (7, 4);

SELECT COALESCE_AGG(num) OVER (ORDER BY id) AS run_start FROM temp;

DROP TABLE temp;
coalesce_agg_result
```

Output:
```
run_start
0          1
1          2
2          2
3          2
4          3
5          3
6          4
```

#### Forward Pass


```python
-- result_4ci 

CREATE OR REPLACE VIEW forward AS


    WITH pre_view AS (
    SELECT
    
    id,
    
    time, 

    value,
    
    RANK() OVER (ORDER BY id) AS run
        
    FROM complete_provided
    
    )
        
    SELECT 
    id, 
    
    time, 

    value, 
    
    run,
    
    COALESCE_AGG(value) OVER (ORDER BY run) AS run_start,
    
    LEAD(value, 1) OVER (ORDER BY run) AS next_val 
    
    FROM pre_view;

-- END SOLUTION

SELECT * FROM forward WHERE next_val IS NOT NULL ORDER BY run LIMIT 100;
```

```python
result_4ci = grading_util.run_file("4ci")[1]
result_4ci
```

Output:
```
id                      time  \
0   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:00:09+00:00   
1   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 11:00:09+00:00   
2   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 11:30:09+00:00   
3   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 11:45:10+00:00   
4   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 12:00:09+00:00   
..                                   ...                       ...   
95  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 10:45:09+00:00   
96  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 11:00:09+00:00   
97  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 11:15:09+00:00   
98  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 11:30:09+00:00   
99  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 11:45:10+00:00   

           value  run     run_start      next_val  
0   65090.004847    1  65095.063848  65155.239319  
1   65155.239319    1  65095.063848  65154.187017  
2   65154.187017    1  65095.063848  6515
```

#### Backward Pass


```python
-- result_4cii
CREATE OR REPLACE VIEW backward AS
-- BEGIN SOLUTION
  WITH pre_view AS (
    SELECT
    
    id,
    
    time, 

    value,
    
    RANK() OVER (ORDER BY id) AS run
        
    FROM complete
    
    ),

    forward AS (
    
    SELECT 
    id, 
    
    time, 

    value, 
    
    run,
    
    COALESCE_AGG(value) OVER (ORDER BY run) AS run_start,
    
    LEAD(value, 1) OVER (ORDER BY run) AS next_val 
    
    FROM pre_view
    )

        
    SELECT 
    id, 
    
    time, 

    value, 
    
    run,
    run_start,
     next_val,

    COALESCE_AGG(next_val) OVER (PARTITION BY run ORDER BY id DESC)AS run_end, 

    ROW_NUMBER() OVER (PARTITION BY run ORDER BY id) AS run_rank,

    COUNT(*) OVER (PARTITION BY run) AS run_size

     
    FROM forward;
-- END SOLUTION

SELECT * FROM backward WHERE run_size > 2 ORDER BY id, run, run_rank LIMIT 100;
```

```python

    SELECT *
    FROM complete;

```

Output:
```
id                      time  \
0       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 23:00:09+00:00   
1       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:45:09+00:00   
2       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:30:09+00:00   
3       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:15:10+00:00   
4       a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:00:09+00:00   
...                                      ...                       ...   
102381  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 13:45:00+00:00   
102382  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 14:00:00+00:00   
102383  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 14:15:00+00:00   
102384  a8d640b1-c17f-51cf-b06f-de6e704e87c1 2018-06-13 09:00:00+00:00   
102385  aa6d2954-2a1f-5a86-9bb2-69ad7326789e 2018-06-13 09:00:00+00:00   

           value  
0       65182.74  
1       65182.51  
2       65182.27  
3       65182.02  
4       65181.77  
...          ...  
102381       NaN  
1
```

```python
result_4cii = grading_util.run_file("4cii")[1]
result_4cii
```

Output:
```
id                      time     value  \
0   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:45:09+00:00  65182.51   
1   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:30:09+00:00  65182.27   
2   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:15:10+00:00  65182.02   
3   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 22:00:09+00:00  65181.77   
4   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-12 21:45:09+00:00  65181.54   
..                                   ...                       ...       ...   
95  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 23:00:09+00:00  65162.09   
96  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 22:45:09+00:00  65161.86   
97  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 22:30:09+00:00  65161.61   
98  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 22:15:09+00:00  65161.36   
99  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-11 22:00:09+00:00  65161.12   

    run  run_start  next_val  run_end  run_rank  run_size  
0     1   65085.9
```

#### Interpolate



```python
-- result_4ciii

CREATE OR REPLACE VIEW likely_data AS
-- BEGIN SOLUTION
 

WITH pre_view AS (
    SELECT 
    id, 
    time,
    value, 
    SUM(CASE WHEN value IS NULL THEN 0 ELSE 1 END) 
         OVER (ORDER BY id) AS run,
         
 COALESCE_AGG(value) OVER (ORDER BY id) AS run_start,
 
    CASE WHEN value IS NULL 
        THEN lead(value, 1) OVER (ORDER BY id)
            ELSE NULL
             END AS next_val

             FROM  complete_provided

     ), 

    prev_two AS (

    SELECT 
    id, 
    time,
    value, 
    run, 
    run_start, 
    next_val, 
    
   CASE WHEN id IS NOT NULL THEN value
            ELSE COALESCE_AGG(next_val) OVER
            (PARTITION BY run ORDER BY id DESC)
        END AS run_end, 

    ROW_NUMBER() OVER (PARTITION BY run ORDER BY id) AS run_rank,
    
    COUNT(*) OVER (PARTITION BY run) AS run_size
    

    FROM pre_view
     
     )
    SELECT 
    id, 
    time,
    value, 
    run, 
    run_start, 
    next_val, 
    run_end, 

    run_rank,
    
    run_size,
    
    run_start + (run_rank - 1)*((run_end-run_start)/(run_size - 1))
         AS interpolated
         
    FROM prev_two;  

    
    
-- END SOLUTION

SELECT * FROM likely_data WHERE run_size > 2 ORDER BY id, time LIMIT 100;
```

```python
SELECT * 
FROM complete_provided;
    

```

Output:
```
id                      time  \
0       a3d47b1a-985e-5395-a6ee-719dad9b580f 2018-06-13 09:00:00+00:00   
1       a3e8e405-0eed-59b6-8747-fe892a6f93de 2018-06-09 04:00:00+00:00   
2       a46fb790-028c-5d17-a3e6-7d08daec0c03 2018-06-13 09:00:00+00:00   
3       a470ac6d-d448-522b-9b74-8ed27f24de25 2018-06-09 04:00:00+00:00   
4       a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 00:30:00+00:00   
...                                      ...                       ...   
102381  a8dc3e90-6bb7-56e4-bf10-bf07694da31b 2018-06-11 02:15:00+00:00   
102382  a8dc3e90-6bb7-56e4-bf10-bf07694da31b 2018-06-11 02:00:00+00:00   
102383  a8dc3e90-6bb7-56e4-bf10-bf07694da31b 2018-06-11 01:45:00+00:00   
102384  a8dc3e90-6bb7-56e4-bf10-bf07694da31b 2018-06-11 01:30:00+00:00   
102385  a8dc3e90-6bb7-56e4-bf10-bf07694da31b 2018-06-11 01:15:00+00:00   

            value  
0             NaN  
1             NaN  
2             NaN  
3             NaN  
4             NaN  
...           ...  
102381  10.32
```

```python
result_4ciii = grading_util.run_file("4ciii")[1]
result_4ciii
```

Output:
```
id                      time  \
0   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:00:09+00:00   
1   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:15:09+00:00   
2   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:30:09+00:00   
3   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:45:09+00:00   
4   a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 01:00:09+00:00   
..                                   ...                       ...   
95  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 23:45:09+00:00   
96  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:00:09+00:00   
97  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:15:09+00:00   
98  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:30:09+00:00   
99  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:45:09+00:00   

           value  run     run_start next_val       run_end  run_rank  \
0   65090.004847  573  65095.063848     None  65090.004847         1   
1   65095.063848  573  65095.063848     None  65095.0
```

Tests for q1b (points: 1)
```
>>> result_1b.shape
(5, 2)
```
```
>>> str(list(result_1b.columns))
"['building', 'json_agg']"
```
```
>>> result_1b['building'].iloc[0]
'ALUMNI HOUSE'
```
```
>>> str(result_1b['json_agg'].iloc[0])
"[{'location': 'BERKELEY', 'building': '1215', 'building_name': 'ALUMNI HOUSE', 'address': 'CORE C22A2M0PUS', 'city_name': 'BERKELEY', 'county': 'ALAMEDA', 'category': 'GENERAL', 'osfg': 15590, 'bldg_asf': 8719, 'levels': 2, 'year': '1954', 'owner': 'UC', 'mp_code': 'P', 'book_value': 106819}, {'location': 'IRVINE', 'building': '9207', 'building_name': 'ALUMNI HOUSE', 'address': 'CORE CAMPUS', 'city_name': 'IRVINE', 'county': 'ORANGE', 'category': 'GENERAL', 'osfg': 4027, 'bldg_asf': 2549, 'levels': 1, 'year': '1984', 'owner': 'UC', 'mp_code': 'P', 'book_value': 41981}, {'location': 'FRANCISC SOAN', 'building': '2032', 'building_name': 'ALUMNI HOUSE', 'address': 'PARNASS US AVE', 'city_name': 'FRANCISC OSAN', 'county': 'FRANCISC OSAN', 'category': 'HEALTH SCIENCE', 'osfg': 7217, 'bldg_asf': 5079, 'levels': 3, 'year': '1915', 'owner': 'UP', 'mp_code': 'P', 'book_value': 135923}]"
```


Tests for q1c (points: 1)
```
>>> result_1c.shape
(4, 2)
```
```
>>> str(list(result_1c.columns))
"['building_name', 'json_agg']"
```
```
>>> result_1c['building_name'].iloc[0]
'EDWARDS FLD'
```
```
>>> str(sorted(result_1c['json_agg'][0], key=lambda x: x['site']))
"[{'site': 'Edwards Stadium East', 'building': 'EDWARDS FLD'}, {'site': 'Edwards Stadium West', 'building': 'EDWARDS FLD'}]"
```


Tests for q2a (points: 3)
```
>>> result_2a.shape
(100, 6)
```
```
>>> str(list(result_2a.columns))
"['time', 'id', 'value', 'median', 'mad', 'is_outlier']"
```
```
>>> result_2a.iloc[:5, :2]
                       time                                    id
0 2018-06-07 00:00:00+00:00  a481d4a8-44f6-5f63-b469-85571a018c6f
1 2018-06-07 00:00:00+00:00  a51d3bed-a8b6-5321-96a1-c5542caca70c
2 2018-06-07 00:00:00+00:00  a6336afc-cdb7-5f87-a4bf-67411ea1b5cc
3 2018-06-07 00:00:00+00:00  a652779f-8433-5ad3-9310-f4f1062e4d04
4 2018-06-07 00:00:00+00:00  a760a23a-8de4-5ce3-939a-826bf31f01df
```
```
>>> result_2a.iloc[:5, 2:]
   value   median     mad  is_outlier
0  812.6  727.500  16.200        True
1   75.0    6.186   4.957        True
2   54.0   32.000   1.000        True
3   14.0    8.000   1.000        True
4  109.0   69.000   7.000        True
```
```
>>> result_2a.iloc[95:, :2]
                        time                                    id
95 2018-06-07 00:00:00+00:00  a7145dcc-90e5-5bfb-9b1c-277615779c16
96 2018-06-07 00:00:00+00:00  a71d821d-d6e6-5ee9-86fa-2b075f13c2cc
97 2018-06-07 00:00:00+00:00  a71fbdfd-26ad-562c-b193-9dc414b851f0
98 2018-06-07 00:00:00+00:00  a744b292-f4cf-5f85-9e4a-5e626e19d5af
99 2018-06-07 00:00:00+00:00  a749b51c-0be8-5b73-bf67-de452e0270ab
```
```
>>> result_2a.iloc[95:, 2:]
    value  median   mad  is_outlier
95    4.0     3.0   1.0       False
96   99.0    81.0  17.0       False
97  122.9    58.7  35.7       False
98  107.0    74.0  18.0       False
99   64.9    64.9   0.0       False
```
```
>>> np.isclose(result_2a.iloc[::2]['value'].sum(), 125485.3427)
True
```
```
>>> np.isclose(result_2a.iloc[::2]['median'].sum(), 122438.12610000001)
True
```
```
>>> np.isclose(result_2a.iloc[::2]['mad'].sum(), 607.3138999999999)
True
```
```
>>> result_2a['is_outlier'].sum()
50
```


Tests for q2b (points: 1)
```
>>> result_2b.shape
(100, 7)
```
```
>>> str(list(result_2b.columns))
"['time', 'id', 'value', 'median', 'mad', 'is_outlier', 'clean_value']"
```
```
>>> result_2b['is_outlier'].sum()
50
```
```
>>> result_2b.iloc[:5, :2]
                       time                                    id
0 2018-06-07 00:00:00+00:00  a481d4a8-44f6-5f63-b469-85571a018c6f
1 2018-06-07 00:00:00+00:00  a51d3bed-a8b6-5321-96a1-c5542caca70c
2 2018-06-07 00:00:00+00:00  a6336afc-cdb7-5f87-a4bf-67411ea1b5cc
3 2018-06-07 00:00:00+00:00  a652779f-8433-5ad3-9310-f4f1062e4d04
4 2018-06-07 00:00:00+00:00  a760a23a-8de4-5ce3-939a-826bf31f01df
```
```
>>> result_2b.iloc[:5, 2:]
   value   median     mad  is_outlier  clean_value
0  812.6  727.500  16.200        True   799.554360
1   75.0    6.186   4.957        True    28.233745
2   54.0   32.000   1.000        True    36.447800
3   14.0    8.000   1.000        True    12.447800
4  109.0   69.000   7.000        True   100.134600
```
```
>>> result_2b.iloc[95:, :2]
                        time                                    id
95 2018-06-07 00:00:00+00:00  a7145dcc-90e5-5bfb-9b1c-277615779c16
96 2018-06-07 00:00:00+00:00  a71d821d-d6e6-5ee9-86fa-2b075f13c2cc
97 2018-06-07 00:00:00+00:00  a71fbdfd-26ad-562c-b193-9dc414b851f0
98 2018-06-07 00:00:00+00:00  a744b292-f4cf-5f85-9e4a-5e626e19d5af
99 2018-06-07 00:00:00+00:00  a749b51c-0be8-5b73-bf67-de452e0270ab
```
```
>>> result_2b.iloc[95:, 2:]
    value  median   mad  is_outlier  clean_value
95    4.0     3.0   1.0       False          4.0
96   99.0    81.0  17.0       False         99.0
97  122.9    58.7  35.7       False        122.9
98  107.0    74.0  18.0       False        107.0
99   64.9    64.9   0.0       False         64.9
```
```
>>> np.isclose(result_2b.iloc[::2]['clean_value'].sum(), 123948.63374060001)
True
```


Tests for q3b (points: 1)
```
>>> result_3b.shape
(1, 2)
```
```
>>> str(list(result_3b.columns))
"['num_unique_units', 'num_unique_units_ignore_case']"
```
```
>>> 0 <= result_3b['num_unique_units'].loc[0] <= 50
True
```
```
>>> 0 <= result_3b['num_unique_units_ignore_case'].loc[0] <= 50
True
```


Tests for q3c (points: 1)
```
>>> result_3c.shape
(1, 1)
```
```
>>> str(list(result_3c.columns))
"['are_units_consistent']"
```
```
>>> result_3c['are_units_consistent'].loc[0] in [True, False]
True
```


Tests for q3e (points: 2)
```
>>> result_3e.shape
(5276, 4)
```
```
>>> str(list(result_3e.columns))
"['building_name', 'address', 'location', 'clean_location']"
```
```
>>> len(result_3e.loc[result_3e['location'] == 'BERKELEY'])
612
```
```
>>> all(result_3e.loc[result_3e['location'] == 'BERKELEY']['clean_location'] == 'BERKELEY')
True
```
```
>>> len(result_3e.loc[result_3e['location'] == 'DAVIS'])
1136
```
```
>>> all(result_3e.loc[result_3e['location'] == 'DAVIS']['clean_location'] == 'DAVIS')
True
```
```
>>> len(result_3e.loc[result_3e['location'] == 'SYSTEMWI DE'])
6
```
```
>>> all(result_3e.loc[result_3e['location'] == 'SYSTEMWI DE']['clean_location'] == 'SYSTEMWIDE')
True
```
```
>>> len(result_3e.loc[result_3e['location'] == 'SAN DSAIENG O'])
1
```
```
>>> all(result_3e.loc[result_3e['location'] == 'SAN DSAIENG O']['clean_location'] == 'SAN DIEGO')
True
```
```
>>> len(result_3e.loc[result_3e['location'] == 'FRANCISC O'])
7
```
```
>>> all(result_3e.loc[result_3e['location'] == 'FRANCISC O']['clean_location'] == 'SAN FRANCISCO')
True
```
```
>>> len(result_3e.loc[result_3e['location'] == result_3e['clean_location']])
5156
```
```
>>> len(result_3e.loc[result_3e['location'] != result_3e['clean_location']])
120
```


Tests for q4a (points: 2)
```
>>> result_4a.shape
(771, 6)
```
```
>>> str(list(result_4a.columns))
"['id', 'time', 'value', 'lag_time', 'lag_value', 'time_diff']"
```
```
>>> result_4a.iloc[:10, :2]
                                     id                      time
0  a3d47b1a-985e-5395-a6ee-719dad9b580f 2018-06-13 09:15:00+00:00
1  a3e8e405-0eed-59b6-8747-fe892a6f93de 2018-06-09 04:15:00+00:00
2  a46fb790-028c-5d17-a3e6-7d08daec0c03 2018-06-13 09:15:00+00:00
3  a470ac6d-d448-522b-9b74-8ed27f24de25 2018-06-09 04:15:00+00:00
4  a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 00:45:00+00:00
5  a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 01:30:00+00:00
6  a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 17:45:00+00:00
7  a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 18:45:00+00:00
8  a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 19:15:00+00:00
9  a48e47ab-bbcd-5b39-aaaa-47fbaafe166e 2018-06-07 19:45:00+00:00
```
```
>>> result_4a.iloc[:10, 2:]
     value                  lag_time  lag_value       time_diff
0    0.056 2018-06-13 08:45:00+00:00      0.056 0 days 00:30:00
1  280.800 2018-06-09 03:45:00+00:00    281.600 0 days 00:30:00
2  281.460 2018-06-13 08:45:00+00:00    281.390 0 days 00:30:00
3    0.004 2018-06-09 03:45:00+00:00      0.006 0 days 00:30:00
4   27.273 2018-06-07 00:15:00+00:00     30.000 0 days 00:30:00
5   27.273 2018-06-07 01:00:00+00:00     23.077 0 days 00:30:00
6   21.429 2018-06-07 17:15:00+00:00     25.000 0 days 00:30:00
7   25.000 2018-06-07 18:15:00+00:00     30.000 0 days 00:30:00
8   20.000 2018-06-07 18:45:00+00:00     25.000 0 days 00:30:00
9   23.077 2018-06-07 19:15:00+00:00     20.000 0 days 00:30:00
```
```
>>> result_4a.iloc[-10:, :2]
                                       id                      time
761  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-12 10:00:00+00:00
762  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-12 10:30:00+00:00
763  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-12 14:00:00+00:00
764  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 05:15:00+00:00
765  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 06:00:00+00:00
766  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 08:15:00+00:00
767  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 13:00:00+00:00
768  a842919a-dcd1-5d91-b1df-f70c988cabfe 2018-06-13 14:30:00+00:00
769  a8d640b1-c17f-51cf-b06f-de6e704e87c1 2018-06-13 09:15:00+00:00
770  aa6d2954-2a1f-5a86-9bb2-69ad7326789e 2018-06-13 09:15:00+00:00
```
```
>>> result_4a.iloc[-10:, 2:]
     value                  lag_time  lag_value       time_diff
761    5.4 2018-06-12 09:00:00+00:00      2.700 0 days 01:00:00
762    2.7 2018-06-12 10:00:00+00:00      5.400 0 days 00:30:00
763    5.4 2018-06-12 10:30:00+00:00      2.700 0 days 03:30:00
764    5.4 2018-06-13 04:15:00+00:00      5.400 0 days 01:00:00
765    2.7 2018-06-13 05:30:00+00:00      5.400 0 days 00:30:00
766    5.4 2018-06-13 07:45:00+00:00      0.012 0 days 00:30:00
767    5.4 2018-06-13 08:15:00+00:00      5.400 0 days 04:45:00
768    2.7 2018-06-13 13:00:00+00:00      5.400 0 days 01:30:00
769  106.0 2018-06-13 08:45:00+00:00    106.000 0 days 00:30:00
770  286.7 2018-06-13 08:45:00+00:00    286.700 0 days 00:30:00
```


Tests for q4b (points: 2)
```
>>> result_4b.shape # shape of first 100 rows in view
(100, 3)
```
```
>>> str(list(result_4b.columns))
"['id', 'time', 'value']"
```
```
>>> result_4b.iloc[:5]
                                     id                      time     value
0  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:00:09+00:00  65085.99
1  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:15:09+00:00  65086.16
2  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:30:09+00:00  65086.35
3  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 00:45:09+00:00  65086.52
4  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 01:00:09+00:00  65086.71
```
```
>>> result_4b.iloc[-5:]
                                      id                      time     value
95  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-07 23:45:09+00:00  65100.44
96  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:00:09+00:00  65100.63
97  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:15:09+00:00  65100.80
98  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:30:09+00:00  65100.97
99  a3d3326f-20ab-5f1d-97c7-f3084df43f06 2018-06-08 00:45:09+00:00  65101.16
```
```
>>> np.isclose(result_4b.iloc[:50]['value'].sum(), 3254484.79)
True
```


Tests for q4ci (points: 1)
```
>>> result_4ci.shape
(100, 6)
```
```
>>> str(list(result_4ci.columns))
"['id', 'time', 'value', 'run', 'run_start', 'next_val']"
```


Tests for q4cii (points: 1)
```
>>> result_4cii.shape
(100, 9)
```
```
>>> str(list(result_4cii.columns))
"['id', 'time', 'value', 'run', 'run_start', 'next_val', 'run_end', 'run_rank', 'run_size']"
```


Tests for q4ciii (points: 1)
```
>>> result_4ciii.shape
(100, 10)
```
```
>>> str(list(result_4ciii.columns))
"['id', 'time', 'value', 'run', 'run_start', 'next_val', 'run_end', 'run_rank', 'run_size', 'interpolated']"
```
```
>>> complete_1 = np.isclose(result_4ciii.iloc[:50]['interpolated'].sum(), 195.94525) # using complete
>>> complete_provided_1 = np.isclose(result_4ciii.iloc[:50]['interpolated'].sum(), 369.7514880904479) # using complete_provided
>>> complete_1 or complete_provided_1
True
```


Tests for q5a (points: 1)
```
>>> result_5a.shape
(1, 1)
```
```
>>> str(list(result_5a.columns))
"['cycle_exists']"
```
```
>>> result_5a['cycle_exists'].iloc[0] in [True, False]
True
```


Tests for q5b (points: 1)
```
>>> result_5b.shape
(1, 1)
```
```
>>> str(list(result_5b.columns))
"['is_tree']"
```
```
>>> result_5b['is_tree'].iloc[0] in [True, False]
True
```


Tests for q5c (points: 3)
```
>>> result_5c.shape
(2, 2)
```
```
>>> str(list(result_5c.columns))
"['sensor_child', 'count']"
```
```
>>> result_5c['sensor_child']
0    https://brickschema.org/schema/Brick#Demand_Sensor
1     https://brickschema.org/schema/Brick#Power_Sensor
Name: sensor_child, dtype: object
```
```
>>> 0 <= result_5c.iloc[0, 1] <= 100
True
```
```
>>> 3100 <= result_5c.iloc[1, 1] <= 3200
True
```
